# What are the most serious meteorological disasters in the United States?

![A photo of a tornado](data/images/tornado.jpg)

## Introduction

**Business Context.** While natural *events* often cannot be avoided, the risks they present can be managed, either by mitigation, avoidance, or insurance, in order to prevent them from becoming natural *disasters*. The consultancy firm you work for has been hired by an independent advocacy group that wants to conduct an analysis of the US emergency management system of preparedness, protection, mitigation, response, and recovery, with the purpose of proposing legislative reforms to make it more effective and financially efficient. Their ultimate goal is to help increase the government's ability to prevent disasters from happening and reduce the negative impact of those that cannot be completely avoided.

**Business Problem.** Your client would like to know which storm event types are more likely to become disasters, and in which locations, as measured by the number of deaths, injuries, and economic damage they cause. Additionally, they would like to conduct a preliminary assessment of whether the [Post-Katrina Emergency Management Reform Act of 2006](https://www.congress.gov/bill/109th-congress/senate-bill/3721) had any impact on the severity of the disasters that occurred after the bill was signed. This Act centralized the US emergency management under the coordination of the Federal Emergency Management Agency (FEMA) as a response to the enormous human and material losses that were caused by Hurricane Katrina in August 2005.

**Analytical Context.** You have been provided with a compressed GZIP file of storm events from 1970 to 2020 as recorded by the US [National Oceanic and Atmospheric Administration](https://www.ncdc.noaa.gov/stormevents/ftp.jsp). You can check the [documentation](https://www1.ncdc.noaa.gov/pub/data/swdi/stormevents/csvfiles/Storm-Data-Bulk-csv-Format.pdf) for more information.

```python
# Importing relevant libraries
import pandas as pd
import matplotlib.pyplot as plt
import seaborn as sns
import folium
import pingouin as pg
```

## Loading in the dataset

Let's load in the dataset. We add the `parse_dates` argument to tell pandas which columns should be interpreted as dates.

```python
df = pd.read_csv("data/dataset.csv.gz", parse_dates=["BEGIN_DATE_TIME", "END_DATE_TIME"])
df.head()
```

Here is a description of the imported columns:

1. **EPISODE_ID**: The storm episode ID. A single episode can contain multiple events
2. **EVENT_ID**: This is the ID of the actual storm as such. Several storms can be grouped into an episode
3. **STATE**: The state or region where the event occurred
4. **EVENT_TYPE**: The type of the event
5. **BEGIN_DATE_TIME**: The date and time when the event started. Times and dates are in LST (Local Solar Time), which means that they reflect the local time, not a coordinated time
6. **BEGIN_YEAR**: The year in which the event begun
7. **CZ_TIMEZONE**: The timezone of the place where the event occurred
8. **END_DATE_TIME**: The date and time when the event ended. Times and dates are in LST (Local Solar Time), which means that they reflect the local time, not a coordinated time
9. **TOR_F_SCALE**: The [enhanced Fujita scale](https://en.wikipedia.org/wiki/Enhanced_Fujita_scale) (highest recorded value). This scale measures the strength of a tornado based on the amount of damage that it caused. A level of `EF0` means "light damage" (wind speeds of 40 - 72 mph), and a level of `EF5` means "incredible damage" (261 - 318 mph). `EFU` means "Unknown"
10. **BEGIN_LOCATION**: The name of the city or village where the event started
11. **END_LOCATION**: The name of the city or village where the event ended
12. **BEGIN_LAT**: The latitude of the place where the event begun
13. **BEGIN_LON**: The longitude of the place where the event begun
14. **END_LAT**: The latitude of the place where the event ended
15. **END_LON**: The longitude of the place where the event ended
16. **TOTAL_DEATHS**: Deaths directly or indirectly attributable to the event
17. **TOTAL_INJURIES**: Injuries directly or indirectly attributable to the event
18.  **TOTAL_DAMAGE_DEFLATED**: Estimated damage to property and crops in dollars. These dollars are "real" dollars, which means that the damages for all the years have been converted ([deflated](https://faculty.fuqua.duke.edu/~rnau/Decision411_2007/411infla.htm)) to the value they would have had in 1982-84. This was done to make the damages comparable across years, since dollars [change purchasing power every year](https://www.insider.com/fast-food-burgers-cost-every-year-2018-9) due to inflation. The deflation was done using the Bureau of Labor Statistics Urban Consumer Price Index, whose base period is 1982-84.

## Exploratory data analysis (EDA)

### Exercise 1

A good first step in EDA is visualizing the univariate distributions of each feature in the dataset. Here's a numbered list of the columns:

| Number 	| Column name 	|
|-	|-	|
| 1 	| STATE 	|
| 2 	| EVENT_TYPE 	|
| 3 	| BEGIN_DATE_TIME 	|
| 4 	| BEGIN_YEAR 	|
| 5 	| CZ_TIMEZONE 	|
| 6 	| END_DATE_TIME 	|
| 7 	| TOR_F_SCALE 	|
| 8 	| BEGIN_LOCATION 	|
| 9 	| END_LOCATION 	|
| 10 	| BEGIN_LAT 	|
| 11 	| BEGIN_LON 	|
| 12 	| END_LAT 	|
| 13 	| END_LON 	|
| 14 	| TOTAL_DEATHS 	|
| 15 	| TOTAL_INJURIES 	|
| 16 	| TOTAL_DAMAGE_DEFLATED 	|

Modify the function below to indicate which visualizations are well-suited to exploring the univariate distributions of each one of the columns. Append the number of the feature to the corresponding list.

Some things to keep in mind:

1. Please provide visualization choices for *each* of the features.
2. If you believe that a visualization doesn't apply to any of the columns, simply leave it as an empty list.
3. A single feature can be added to more than one visualization's corresponding list.

<!-- grade_id: cell-44a893d8ae368d56, points: 0, grade: false, solution: true -->
```python
def visualization_choice():
    """
    Returns the answer to exercise 1.
    
    Example:
    Let's say you want to add features 1 and 16 to `barplot`. Just
    add the numbers to the list. In other words, the line should
    be `barplot = [1,16]`.
    """
    
    # HISTOGRAMS
    histogram = []
    
    # BAR PLOTS
    barplot = []
    
    # BOX PLOTS
    boxplot = []
    
    # GEO HEAT MAP
    geoheatmap = []
    
    # LINE PLOT
    lineplot = []
    
    # DENSITY PLOT
    densityplot = []
    
    # YOUR CODE HERE
    raise NotImplementedError() # Remove this line when you enter your solution

    
    visualization_choices = {
        "histogram":histogram,
        "barplot":barplot,
        "boxplot":boxplot,
        "geoheatmap":geoheatmap,
        "lineplot":lineplot,
        "densityplot":densityplot,
    }
    
    return visualization_choices

visualization_choice()
```

### Exercise 2

#### 2.1

As discussed in the regular cases, it is important to compute summary statistics for each feature. We know that the damage caused by natural disasters usually has a distribution that is very skewed to the right. Which TWO of the following summary statistics would be MOST useful to confirm that this is the case for the `TOTAL_DEATHS`, `TOTAL_INJURIES`, and `TOTAL_DAMAGE_DEFLATED` columns?
<ul>
A. Count<br>
B. Mean<br>
C. Minimum<br>
D. Maximum<br>
E. 25th percentile<br>
F. Median<br>
G. 75th percentile<br>
</ul>

<!-- grade_id: cell-6cc37984b8bd9b04, points: 0, grade: false, solution: true -->
```python
def summary_choice():
    """
    Returns the two options you chose.
    
    Uncomment the lines that correspond to your choice.
    """
    
    A = df.groupby(["EVENT_TYPE"])[["TOTAL_DEATHS", "TOTAL_INJURIES", "TOTAL_DAMAGE_DEFLATED"]].count()
    B = df.groupby(["EVENT_TYPE"])[["TOTAL_DEATHS", "TOTAL_INJURIES", "TOTAL_DAMAGE_DEFLATED"]].mean()
    C = df.groupby(["EVENT_TYPE"])[["TOTAL_DEATHS", "TOTAL_INJURIES", "TOTAL_DAMAGE_DEFLATED"]].min()
    D = df.groupby(["EVENT_TYPE"])[["TOTAL_DEATHS", "TOTAL_INJURIES", "TOTAL_DAMAGE_DEFLATED"]].max()
    E = df.groupby(["EVENT_TYPE"])[["TOTAL_DEATHS", "TOTAL_INJURIES", "TOTAL_DAMAGE_DEFLATED"]].quantile(0.25)
    F = df.groupby(["EVENT_TYPE"])[["TOTAL_DEATHS", "TOTAL_INJURIES", "TOTAL_DAMAGE_DEFLATED"]].quantile(0.5)
    G = df.groupby(["EVENT_TYPE"])[["TOTAL_DEATHS", "TOTAL_INJURIES", "TOTAL_DAMAGE_DEFLATED"]].quantile(0.75)
    
    
    answer = {}
    #answer["A"] = A
    #answer["B"] = B
    #answer["C"] = C
    #answer["D"] = D
    #answer["E"] = E
    #answer["F"] = F
    #answer["G"] = G
    
    # YOUR CODE HERE
    raise NotImplementedError() # Remove this line when you enter your solution
    
    return answer
```

#### 2.2

Inspect the tables of the two summary statistics you chose. What can you interpret from your results?

**Note:** In this and all similar subsequent exercises, please answer in the cell that is immediately below this one (if you write in the same cell as the question, your answer will not be recorded).


**Your answer here**.

### Exercise 3


Let's now look at some bivariate distributions. This is a pair plot of some of the numeric variables of the dataset:

![Pairplot](data/images/pairplot.png)

What useful insights can you gather from the plot?


**Your answer here.**

### Exercise 4

Some storm events have a wider risk area than others, which can make them more worthy of attention for authorities. Which of these visualization strategies could be useful to find patterns related to the size of the risk area around storm events? Append the letters of all options that apply to the `answer` list in the function below.

&nbsp;&nbsp;&nbsp;&nbsp; A. A scatter plot of `BEGIN_LAT` vs. `END_LAT`<br>
&nbsp;&nbsp;&nbsp;&nbsp; B. A scatter plot of `BEGIN_LON` vs. `END_LON`<br>
&nbsp;&nbsp;&nbsp;&nbsp; C. A geographic heat map of the count of `BEGIN_LOCATION` per event<br>
&nbsp;&nbsp;&nbsp;&nbsp; D. A grouped geographic scatterplot of those events whose `BEGIN_LOCATION` and `END_LOCATION` do not coincide (in one color) vs. those in which they do coincide (in another color). (A geographic scatterplot is a scatterplot layered on top of a map, with the $x$- and $y$-coordinates being the longitude and the latitude of each point - see some examples [here](https://plotly.com/python/scatter-plots-on-maps/).)

<!-- grade_id: cell-90a2fa1b33bd82c5, points: 0, grade: false, solution: true -->
```python
def geo_spread():
    """
    Returns the selected answer
    """
    
    answer = []
    
    # YOUR CODE HERE
    raise NotImplementedError() # Remove this line when you enter your solution
    
    # answer.append(["A"])
    # answer.append(["B"])
    # answer.append(["C"])
    # answer.append(["D"])
    
    return answer
```

### Exercise 5

#### 5.1

Plot one (or more) of your choices from the previous exercise.

<!-- grade_id: cell-35f481a6ddbc527d, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError() # Remove this line when you enter your solution
```

#### 5.2

Write down your findings in the cell below.


**Your answer here.**

### Exercise 6

In the context of natural disaster management, risk is defined as the probability of an event happening times the hazard that its occurrence may pose (that is, the potential damages that could happen if the event takes place):

$$Risk = Probability \times Hazard$$

The estimation of risk is part of what is sometimes called "risk assessment". 

Since we are still in the exploratory phase, we do not want to start modeling right away. We would like to rely on plots and visualizations instead.

Describe a data visualization strategy that includes metrics and plots of different kinds to estimate:

1. which storm types are most likely to happen in a given month;
2. how large the economic damages caused by the storms would be; and
3. in which locations the storms are most likely to happen


**Your answer here.**

### Exercise 7

Code at least one of the visualizations you proposed in the previous exercise.

<!-- grade_id: cell-47731d5b2e81022b, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError() # Remove this line when you enter your solution
```

## The Post-Katrina Emergency Management Reform Act of 2006

### Exercise 8 (hard)

#### 8.1

Conduct a hypothesis test for each event type to assess whether there is a difference in average total damage when comparing disasters that happened before the reform to those that happened after. Keep only the event types for which you found a significant difference (using a significance threshold of $\alpha=0.01$). Since it is likely that not all events that have happened in the US are present in this dataset, we can interpret the data as being a sample (conducting hypothesis tests on population data would not make sense).

**Note:** If an event type does not have associated events either before or after the Act, ignore it and don't add it to the dictionary (since a $t$ - test won't be possible). Do not apply the Bonferroni correction in this exercise.

<details>
    <summary markdown="span">
        <br>Click here for a <b>Hint</b>
    </summary>
    <blockquote>
        Use the <a href="https://pingouin-stats.org/generated/pingouin.ttest.html"><code>.ttest()</code></a> function from the <code>pingouin</code> library.
    </blockquote>
</details>

<!-- grade_id: cell-d8057066d450d1b7, points: 0, grade: false, solution: true -->
```python
df["POST_ACT"] = df["BEGIN_YEAR"] > 2006

def test_differences(df):
    """
    Conducts a t-test on TOTAL_DAMAGES comparing events
    that happened in 2006 or before with events that
    happenned after that year.
    
    Inputs:
    `df`: A pandas DataFrame
    
    Outputs:
    `p_values_signif`: A Python dictionary in which the keys are the event type
    and the values are the significant p-values that resulted from the t-test (alpha
    of 0.01)
    
    Note: If an event type does not have associated events either before or
    after the act, ignore it and don't add it to the dictionary (since a t-test
    won't be possible)
    """
      
    pre_dam = df[df["POST_ACT"]==False][["EVENT_TYPE", "TOTAL_DAMAGE_DEFLATED"]].dropna(how="any")
    post_dam = df[df["POST_ACT"]==True][["EVENT_TYPE", "TOTAL_DAMAGE_DEFLATED"]].dropna(how="any")
    
    # YOUR CODE HERE
    raise NotImplementedError() # Remove this line when you enter your solution
    return p_values_signif
```

#### 8.2

For these significant event types, plot their total deflated damages as box plots, comparing the pre-Act events with the post-Act events.

<!-- grade_id: cell-e7b19a29d928da01, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError() # Remove this line when you enter your solution
```

#### 8.3

What conclusions can you draw from the hypothesis tests you have just conducted, and why?


**Your answer here.**

## Testing cells

<!-- grade_id: cell-edc3ec41ffe69c52, points: 4, grade: true, solution: false -->
```python
# Ex. 1
ex1 = visualization_choice()
c = set()
for key in ex1.keys():
    for i in ex1[key]:
        c.add(i)
assert len(c - {1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16}) == 0, "Ex. 1 - It seems you included a number that is not one of the column IDs!"
assert c == {1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16}, "Ex. 1 - It seems you did not add a visualization for all the columns!"
print("Exercise 1 passed our preliminary screening. You will get the final grade after your TA has reviewed your submission!")
```

<!-- grade_id: cell-bb9ac3e19bfa2b3d, points: 2, grade: true, solution: false -->
```python
# Ex. 2
ans_2 = summary_choice().keys()
assert len(ans_2) == 2, "Ex. 2.1 - Please choose exactly two summary statistics!"
print("Exercise 2.1 passed our preliminary screening. You will get the final grade after your TA has reviewed your submission!")
```

<!-- grade_id: cell-65ffb1c301fcb67b, points: 2, grade: true, solution: false -->
```python
# Ex. 4
assert len(geo_spread()) > 0, "Ex. 4 - Please choose at least one option!"
print("Exercise 4 passed our preliminary screening. You will get the final grade after your TA has reviewed your submission!")
```

<!-- grade_id: cell-f7caf6bc179cc69a, points: 4, grade: true, solution: false -->
```python
# Ex. 8.1
import warnings
warnings.filterwarnings("ignore")
p_values_signif = test_differences(df)
assert len(p_values_signif)==9, "Ex. 8.1 - Your output seems to have the wrong number of elements! Remember to use an alpha of 0.01 (not 0.05)"
s = 0
for key in p_values_signif.keys():
    s = s+p_values_signif[key]
assert s < 0.015, "Ex. 8.1 - Your p-values are apparently too large! Remember to use an alpha of 0.01 (not 0.05)"
print("Exercise 8.1 passed our preliminary screening. You will get the final grade after your TA has reviewed your submission!")
print("Exercise 8.1 looks correct!")
```

## Attribution

"Storm Events Database", 20 Nov 2020, National Oceanic and Atmospheric Administration, Licensed under the [Freedom of Information Act](https://www.noaa.gov/organization/information-technology/freedom-of-information-act), https://www.ncdc.noaa.gov/stormevents/ftp.jsp

"F5 tornado Elie Manitoba 2007", 22 Jun 2017, Justin Hobson, Creative Commons Attribution-Share Alike 3.0 Unported license, https://commons.wikimedia.org/wiki/File:F5_tornado_Elie_Manitoba_2007.jpg

"Historical Consumer Price Index for All Urban Consumers (CPI-U): U.S. city average, all items, index averages", March, 2021, Licensed under the [Freedom of Information Act](https://www.bls.gov/bls/blsfoia.htm), https://www.bls.gov/cpi/tables/supplemental-files/historical-cpi-u-202103.pdf (the value for 2021 corresponds to the three-month average between January and March).